```python
# Initialize Otter
import otter
grader = otter.Notebook("helloWorld.ipynb")
```

Welcome to class!

All assignments will be in this format. The file you are working in is called a Jupyter Notebook. It is slightly different from a regular Python file, but the concepts are the same.

The first 4 assignments are a very quick overview of basic syntax and data structures related to Python. This course assumes you have some programming experience (perhaps: matlab, or R). If you do not, and you find the content moving too fast, there are many resources available online.

1. https://cse.msu.edu/~cse231/ - open lectures and content relating to the undergraduate Python course from the MSU CSE department.
2. https://www.geeksforgeeks.org/python-programming-language-tutorial/?ref=outindfooter - in-depth Python tutorial and resource website
3. https://www.w3schools.com/python/default.asp - another web-based Python tutorial and reference website
4. https://medium.com/ - an open blog site that contains many, many worked examples of different projects in different languages
5. Me, your instructor!

In this lecture, we will cover:
* Jupyter notebooks, including:
    * What they are
    * How to run them
    * How they are different from Python `.py` scripts
* Defining/assigning variables
* Printing variables using `print()`, or displaying variables by calling their name
* Comments, multi-line and in-line

## Lecture Section

You can run each cell by pressing Shift + Enter, or by clicking the "Run" button in the toolbar above (may be next to the code block, depending on your IDE). The first cell is a simple print statement, with a string inside of it. Try running it to see what it does!

```python
print('Hello, World!')
```

One nice thing about notebooks is that you can see the output of your code right below the cell where you wrote it. This makes it easy to see what your code is doing and to debug any errors. You don't even need a print statement to see the output of the last line in a cell.

Try running the next two cells to see what they do!

```python
'Hello, World!'
```

If you aren't using a print() statement, make sure you put the value you want to see on the last line of the cell. Otherwise, you won't see anything when you run the cell.

Run the code below as is, then delete the last line and run it again.


```python
x = 5
x
y = 10.8
```

Above, we have variable `x` and `y`. They are on the lefthand side of the `=` and they hold the value on the right hand side. That is why placing an `x` and removing the last line returns `5` in the output below the code block.

Your variables should always be defined with a number (like 5 and 10 above), a string (such as 'Hello, World!), or a data structure (covered in the future). Numbers are further broken into integers (int) and floating point (floats). Integers are whole numbers, and floating point numbers are decimals.


You can define multiple variables in the same line, too.

```python
x, y, z = 1, 2, 3
print(x, y, z)
```

If you try to output a variable that is undefined, you will get an error.

```python
hello_world
```


A key difference between Jupyter Notebooks and regular Python files is that you can have multiple cells, like you've already seen! Each cell can contain code, and you can run them one at a time. This allows you to test small pieces of code without running the entire file.

Try your own print statement in the next cell!

```python
# put a print statement below
```

Another difference between Jupyter Notebooks and regular Python files is that you can have text cells like this one. You can use these cells to write notes, explanations, or anything else you want.

Try writing a short message in the next cell!

Lastly, if you want to comment your code, simply put a # in the code block and your comment after. If you need to make a long comment, you can use triple-quoted literals.

```python
'''
I am a multi-line comment.
I can have multiple lines of text.
'''

# I am just an in-line comment
```

## Assignment Section

**Question 1.** For each of the variables below, assign strings or numbers to answer the questions.


```python
name = ...
# favorite WHOLE number...
favorite_number = ...
favorite_color = ...
favorite_food = ...
favorite_month = ...
favorite_decimal = ...
favorite_hobby = ...

```

```python
grader.check("q1")
```

**Question 2.** Find the Euclidean distance between the two points (3, 4) and (5, 6). To do a sqrt(), use 'math.sqrt()' from the math package. I already imported it for you, and I put it where it needs to go - the thing you want to take the square-root of should go in the parentheses of 'math.sqrt()'. We will talk about libraries and `import` statements in a future lecture.


```python
import math
euc_dist = math.sqrt(...)

```

```python
grader.check("q2")
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for q1 (points: None)
```
>>> assert isinstance(name, str)
>>> assert isinstance(favorite_color, str)
>>> assert isinstance(favorite_food, str)
>>> assert isinstance(favorite_month, str)
>>> assert isinstance(favorite_hobby, str)
```
```
>>> assert isinstance(favorite_number, int)
```
```
>>> assert isinstance(favorite_decimal, float)
```


Tests for q2 (points: None)
```
>>> math.isclose(euc_dist, 2.82)
False
```
